---
title: 'Essential Packages'
description: 'Install the essential packages for sensory data analysis'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise.asasasasasaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa

`@instructions`
Install adahdal skhfshfss;lfsfsf

`@hint`
kahdkhdiehkblkbsklhdkhsfh

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
install.packages(____)
```

`@solution`
```{r}

```

`@sct`
```{r}

```
